---
$schema: /app-interface/prometheus-rule-test-1.yml

rule_files:
  - /observability/prometheusrules/uhc-acct-mngr.prometheusrules.yaml

evaluation_interval: 1m

tests:
  - interval: 1h
    input_series:
      # banned_users
      - series: banned_users{namespace="uhc-{{{environment}}}",service="uhc-acct-mngr-metrics",pod="uhc-acct-mngr-7fccd765b-nscfn"}
        values: 0+2x24 # 2 banned users per hour
      - series: banned_users{namespace="uhc-{{{environment}}}",service="uhc-acct-mngr-metrics",pod="uhc-acct-mngr-7fccd765b-9nz6b"}
        values: 0+2x24 # duplicate series
    alert_rule_test:
      - eval_time: 12h
        alertname: UHC account manager too many banned users - {{{environment}}}
        exp_alerts:
      - eval_time: 24h
        alertname: UHC account manager too many banned users - {{{environment}}}
        exp_alerts:
          - exp_labels:
              namespace: uhc-{{{environment}}}
              service: uhc-acct-mngr
              severity: medium
            exp_annotations:
              message: "UHC Account Manager in namespace uhc-{{{environment}}} has banned 48 users in 24 hours"
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/uhc/sop"

  - interval: 5m
    input_series:
    - series: certificate_lifespan{namespace="uhc-{{{environment}}}",service="uhc-acct-mngr-metrics",type="rhit"}
      values: 30 30 30 29 29
    alert_rule_test:
      - eval_time: 10m
        alertname: AMS RHIT certificate expiring in less than 30 days - {{{environment}}}
        exp_alerts: []
      - eval_time: 20m
        alertname: AMS RHIT certificate expiring in less than 30 days - {{{environment}}}
        exp_alerts:
        - exp_labels:
            type: rhit
            namespace: uhc-{{{environment}}}
            service: uhc-acct-mngr
            severity: warning
          exp_annotations:
            runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/uhc/sop#account-manager-client-certificate-expiring"
            message: "AMS RHIT certificate will expire in less than 30 days - {{{environment}}}"

  - interval: 5m
    input_series:
    - series: certificate_lifespan{namespace="uhc-{{{environment}}}",service="uhc-acct-mngr-metrics",type="rhit"}
      values: 7 7 7 6 6
    alert_rule_test:
      - eval_time: 10m
        alertname: AMS RHIT certificate expiring in less than 7 days - {{{environment}}}
        exp_alerts: []
      - eval_time: 20m
        alertname: AMS RHIT certificate expiring in less than 7 days - {{{environment}}}
        exp_alerts:
        - exp_labels:
            type: rhit
            namespace: uhc-{{{environment}}}
            service: uhc-acct-mngr
            severity: high
          exp_annotations:
            runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/uhc/sop#account-manager-client-certificate-expiring"
            message: "AMS RHIT certificate will expire in less than 7 days - {{{environment}}}"

  - interval: 1m
    input_series:
      # subscription labels
      - series: api_inbound_request_count{code="500",namespace="uhc-{{{environment}}}",service="uhc-acct-mngr-metrics",path="/labels"}
        values: 0+1x20 # http 500 (10%)
      - series: api_inbound_request_count{code="403",namespace="uhc-{{{environment}}}",service="uhc-acct-mngr-metrics",path="/labels"}
        values: 0+1x20 # http 403 (10%)
      - series: api_inbound_request_count{code="200",namespace="uhc-{{{environment}}}",service="uhc-acct-mngr-metrics",path="/labels"}
        values: 0+8x20 # http 200 (80%)
    alert_rule_test:
      - eval_time: 11m
        alertname: UHC account manager subscription labels 5xx Errors High - {{{environment}}}
        exp_alerts:
        - exp_labels:
            service: uhc-acct-mngr
            severity: {{{severity_500_errors}}}
          exp_annotations:
            runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/uhc/sop#subscription-labels"
            message: "UHC Account Manager subscription labels is returning code 5xx for 10% of requests."

  - interval: 1m
    input_series:
      # ams up
      - series: up{job="uhc-acct-mngr-metrics",namespace="uhc-{{{environment}}}",pod="uhc-acct-mngr-7fccd765b-nscfn"}
        values: 1 1 1 1 1 0 0 0 0 0 # up for 5min then down
      # ams inbound request count
      - series: api_inbound_request_count{code="500",namespace="uhc-{{{environment}}}",service="uhc-acct-mngr-metrics"}
        values: 0+1x20 # http 500 (10%)
      - series: api_inbound_request_count{code="403",namespace="uhc-{{{environment}}}",service="uhc-acct-mngr-metrics"}
        values: 0+1x20 # http 403 (10%)
      - series: api_inbound_request_count{code="200",namespace="uhc-{{{environment}}}",service="uhc-acct-mngr-metrics"}
        values: 0+8x20 # http 200 (80%)
      # ams inbound request duration
      - series: api_inbound_request_duration_bucket{le="0.1",namespace="uhc-{{{environment}}}",service="uhc-acct-mngr-metrics"}
        values: 0+1x20 # <= 0.1
      - series: api_inbound_request_duration_bucket{le="1",namespace="uhc-{{{environment}}}",service="uhc-acct-mngr-metrics"}
        values: 0+2x20 # <= 1
      - series: api_inbound_request_duration_bucket{le="2",namespace="uhc-{{{environment}}}",service="uhc-acct-mngr-metrics"}
        values: 0+3x20 # <= 2
      - series: api_inbound_request_duration_bucket{le="5",namespace="uhc-{{{environment}}}",service="uhc-acct-mngr-metrics"}
        values: 0+4x20 # <= 5
      - series: api_inbound_request_duration_bucket{le="10",namespace="uhc-{{{environment}}}",service="uhc-acct-mngr-metrics"}
        values: 0+5x20 # <= 10
      - series: api_inbound_request_duration_bucket{le="30",namespace="uhc-{{{environment}}}",service="uhc-acct-mngr-metrics"}
        values: 0+6x20 # <= 30
      - series: api_inbound_request_duration_bucket{le="+Inf",namespace="uhc-{{{environment}}}",service="uhc-acct-mngr-metrics"}
        values: 0+7x20 # <= inf (50% > 2s)
      #ams advisory lock duration
      - series: advisory_lock_duration_bucket{le="0.1",namespace="uhc-{{{environment}}}",service="uhc-acct-mngr-metrics",type="quota"}
        values: 0+1x20 # <= 0.1
      - series: advisory_lock_duration_bucket{le="0.2",namespace="uhc-{{{environment}}}",service="uhc-acct-mngr-metrics",type="quota"}
        values: 0+2x20 # <= 0.2
      - series: advisory_lock_duration_bucket{le="0.5",namespace="uhc-{{{environment}}}",service="uhc-acct-mngr-metrics",type="quota"}
        values: 0+3x20 # <= 0.5
      - series: advisory_lock_duration_bucket{le="1",namespace="uhc-{{{environment}}}",service="uhc-acct-mngr-metrics",type="quota"}
        values: 0+4x20 # <= 1
      - series: advisory_lock_duration_bucket{le="2",namespace="uhc-{{{environment}}}",service="uhc-acct-mngr-metrics",type="quota"}
        values: 0+5x20 # <= 2
      - series: advisory_lock_duration_bucket{le="10",namespace="uhc-{{{environment}}}",service="uhc-acct-mngr-metrics",type="quota"}
        values: 0+6x20 # <= 10
      - series: advisory_lock_duration_bucket{le="+Inf",namespace="uhc-{{{environment}}}",service="uhc-acct-mngr-metrics",type="quota"}
        values: 0+10x20 # <= inf (50% > 1s)
      # kube_job_status_failed
      - series: kube_job_status_failed{job="kube-state-metrics",namespace="uhc-{{{environment}}}",job_name="uhc-acct-mngr-account-reconciler-1613846160"}
        values: 0+0x59 1+0x59 # 0 for one hr and then 1
      # quay request count
      - series: api_outbound_request_count{apiservice="quay",code="400",namespace="uhc-{{{environment}}}",service="uhc-acct-mngr-metrics"}
        values: 0+1x20 # http 400 10%
      - series: api_outbound_request_count{apiservice="quay",code="0",namespace="uhc-{{{environment}}}",service="uhc-acct-mngr-metrics"}
        values: 0+1x20 # http timeout 10%
      - series: api_outbound_request_count{apiservice="quay",code="200",namespace="uhc-{{{environment}}}",service="uhc-acct-mngr-metrics"}
        values: 0+4x20 # http 200 40%
      - series: api_outbound_request_count{apiservice="quay",code="201",namespace="uhc-{{{environment}}}",service="uhc-acct-mngr-metrics"}
        values: 0+4x20 # http 201 40%
      # quay request duration
      - series: api_outbound_request_duration_bucket{le="1",namespace="uhc-{{{environment}}}",service="uhc-acct-mngr-metrics",apiservice="quay",code="201",path="/api/v1/organization/-/robots/-"}
        values: 0+1x20 # <= 1
      - series: api_outbound_request_duration_bucket{le="5",namespace="uhc-{{{environment}}}",service="uhc-acct-mngr-metrics",apiservice="quay",code="201",path="/api/v1/organization/-/robots/-"}
        values: 0+2x20 # <= 5
      - series: api_outbound_request_duration_bucket{le="10",namespace="uhc-{{{environment}}}",service="uhc-acct-mngr-metrics",apiservice="quay",code="201",path="/api/v1/organization/-/robots/-"}
        values: 0+3x20 # <= 10
      - series: api_outbound_request_duration_bucket{le="+Inf",namespace="uhc-{{{environment}}}",service="uhc-acct-mngr-metrics",apiservice="quay",code="201",path="/api/v1/organization/-/robots/-"}
        values: 0+4x20 # <= inf (50% > 1s)
      # rhit request count
      - series: api_outbound_request_count{apiservice="rhit",code="400",namespace="uhc-{{{environment}}}",service="uhc-acct-mngr-metrics"}
        values: 0+1x20 # http 400 10%
      - series: api_outbound_request_count{apiservice="rhit",code="0",namespace="uhc-{{{environment}}}",service="uhc-acct-mngr-metrics"}
        values: 0+1x20 # http timeout 10%
      - series: api_outbound_request_count{apiservice="rhit",code="200",namespace="uhc-{{{environment}}}",service="uhc-acct-mngr-metrics"}
        values: 0+4x20 # http 200 40%
      - series: api_outbound_request_count{apiservice="rhit",code="201",namespace="uhc-{{{environment}}}",service="uhc-acct-mngr-metrics"}
        values: 0+4x20 # http 201 40%
      # rhit/terms request count
      - series: api_outbound_request_count{apiservice="rhit/terms",code="400",namespace="uhc-{{{environment}}}",service="uhc-acct-mngr-metrics"}
        values: 0+1x20 # http 400 10%
      - series: api_outbound_request_count{apiservice="rhit/terms",code="0",namespace="uhc-{{{environment}}}",service="uhc-acct-mngr-metrics"}
        values: 0+1x20 # http timeout 10%
      - series: api_outbound_request_count{apiservice="rhit/terms",code="200",namespace="uhc-{{{environment}}}",service="uhc-acct-mngr-metrics"}
        values: 0+4x20 # http 200 40%
      - series: api_outbound_request_count{apiservice="rhit/terms",code="201",namespace="uhc-{{{environment}}}",service="uhc-acct-mngr-metrics"}
        values: 0+4x20 # http 201 40%
      # rhit/terms request duration
      - series: api_outbound_request_duration_bucket{le="1",namespace="uhc-{{{environment}}}",service="uhc-acct-mngr-metrics",apiservice="rhit/terms",code="200"}
        values: 0+1x20 # <= 1
      - series: api_outbound_request_duration_bucket{le="5",namespace="uhc-{{{environment}}}",service="uhc-acct-mngr-metrics",apiservice="rhit/terms",code="200"}
        values: 0+2x20 # <= 5
      - series: api_outbound_request_duration_bucket{le="10",namespace="uhc-{{{environment}}}",service="uhc-acct-mngr-metrics",apiservice="rhit/terms",code="200"}
        values: 0+3x20 # <= 10
      - series: api_outbound_request_duration_bucket{le="+Inf",namespace="uhc-{{{environment}}}",service="uhc-acct-mngr-metrics",apiservice="rhit/terms",code="200"}
        values: 0+4x20 # <= inf (50% > 1s)
      # pod restart
      - series: kube_pod_container_status_restarts_total{container="service",namespace="uhc-{{{environment}}}",pod="uhc-acct-mngr-7fccd765b-mpqmk"}
        values: 0 0 1 1 # ams-service
      - series: kube_pod_container_status_restarts_total{container="uhc-acct-mngr-fetch-telemeter-metrics",namespace="uhc-{{{environment}}}",pod="uhc-acct-mngr-fetch-telemeter-metrics-1613509500-srnkx"}
        values: 0 0 2 2 # fetch-telemeter
      # hydra request count
      - series: api_outbound_request_count{apiservice="ocm-/hydra/rest",code="500",namespace="uhc-{{{environment}}}",service="uhc-acct-mngr-metrics"}
        values: 0+1x20 # http 500 10%
      - series: api_outbound_request_count{apiservice="ocm-/hydra/rest",code="0",namespace="uhc-{{{environment}}}",service="uhc-acct-mngr-metrics"}
        values: 0+1x20 # http timeout 10%
      - series: api_outbound_request_count{apiservice="ocm-/hydra/rest",code="200",namespace="uhc-{{{environment}}}",service="uhc-acct-mngr-metrics"}
        values: 0+4x20 # http 200 40%
      - series: api_outbound_request_count{apiservice="ocm-/hydra/rest",code="201",namespace="uhc-{{{environment}}}",service="uhc-acct-mngr-metrics"}
        values: 0+4x20 # http 201 40%
      # hydra request duration
      - series: api_outbound_request_duration_bucket{namespace="uhc-{{{environment}}}",service="uhc-acct-mngr-metrics",apiservice="ocm-/hydra/rest",le="10",code="200"}
        values: 0+90x20 # 90% le=10
      - series: api_outbound_request_duration_count{namespace="uhc-{{{environment}}}",service="uhc-acct-mngr-metrics",apiservice="ocm-/hydra/rest"}
        values: 0+100x20
      # cron job
      - series: job_results_job_count{jobname="ocp-usage-reconciler",namespace="uhc-{{{environment}}}",status="fail"}
        values: 0+0x59 1+0x59 # number of failed jobs is 0 for one hr and then 1
      - series: job_results_job_count{jobname="ocp-usage-reconciler",namespace="uhc-{{{environment}}}",status="pass"}
        values: 1+0x59 0+0x59 # number of passed jobs is 1 for one hr and then 0

    alert_rule_test:
      - eval_time: 4m
        alertname: UHC account manager DOWN - {{{environment}}}
        exp_alerts:
      - eval_time: 10m
        alertname: UHC account manager DOWN - {{{environment}}}
        exp_alerts:
          - exp_labels:
              service: uhc-acct-mngr
              severity: {{{severity_ams_down}}}
            exp_annotations:
              message: "No targets found for UHC Account Manager in namespace uhc-{{{environment}}}. Account Manager is down."
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/uhc/sop"
      - eval_time: 10m
        alertname: UHC account manager 4xx Errors High - {{{environment}}}
        exp_alerts:
      - eval_time: 20m
        alertname: UHC account manager 4xx Errors High - {{{environment}}}
        exp_alerts:
          - exp_labels:
              namespace: uhc-{{{environment}}}
              service: uhc-acct-mngr
              severity: medium
            exp_annotations:
              message: "UHC Account Manager in namespace uhc-{{{environment}}} is returning 4xx errors for 10% of requests."
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/uhc/sop"
      - eval_time: 10m
        alertname: UHC account manager Latency High - {{{environment}}}
        exp_alerts:
      - eval_time: 20m
        alertname: UHC account manager Latency High - {{{environment}}}
        exp_alerts:
          - exp_labels:
              service: uhc-acct-mngr
              severity: medium
            exp_annotations:
              message: "UHC Account Manager in namespace uhc-{{{environment}}} 50th percentile latency is greater than 2s. Current value: 3.5"
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/uhc/sop"
      - eval_time: 10m
        alertname: UHC account manager Advisory Lock Duration High - {{{environment}}}
        exp_alerts:
      - eval_time: 20m
        alertname: UHC account manager Advisory Lock Duration High - {{{environment}}}
        exp_alerts:
          - exp_labels:
              service: uhc-acct-mngr
              severity: medium
            exp_annotations:
              message: "UHC Account Manager Advisory Lock Duration in namespace uhc-{{{environment}}} 50th percentile latency is greater than 1s. Current value: 2"
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/uhc/sop#account-manager-advisory-lock-issues"
      - eval_time: 60m
        alertname: KubeJobFailed
        exp_alerts:
      - eval_time: 120m
        alertname: KubeJobFailed
        exp_alerts:
          - exp_labels:
              job: kube-state-metrics
              job_name: uhc-acct-mngr-account-reconciler-1613846160
              namespace: uhc-{{{environment}}}
              service: uhc-acct-mngr
              severity: medium
            exp_annotations:
              message: "Job uhc-{{{environment}}}/uhc-acct-mngr-account-reconciler-1613846160 failed to complete in uhc-{{{environment}}}."
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/uhc/sop"
      - eval_time: 10m
        alertname: OCM dependencies - Quay too many errors - {{{environment}}}
        exp_alerts:
      - eval_time: 20m
        alertname: OCM dependencies - Quay too many errors - {{{environment}}}
        exp_alerts:
          - exp_labels:
              service: uhc-acct-mngr
              severity: medium
            exp_annotations:
              message: "AMS is seeing a high rate of errors from Quay for 20% of requests."
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/uhc/sop"
      - eval_time: 10m
        alertname: OCM dependencies - Quay creating robot users slow - {{{environment}}}
        exp_alerts:
      - eval_time: 20m
        alertname: OCM dependencies - Quay creating robot users slow - {{{environment}}}
        exp_alerts:
          - exp_labels:
              service: uhc-acct-mngr
              severity: medium
            exp_annotations:
              message: "AMS is seeing a high average latency creating robot accounts on Quay. The average response time over the last 5 minutes is 5 seconds"
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/uhc/sop"
      - eval_time: 10m
        alertname: OCM dependencies - RHIT too many errors - {{{environment}}}
        exp_alerts:
      - eval_time: 20m
        alertname: OCM dependencies - RHIT too many errors - {{{environment}}}
        exp_alerts:
          - exp_labels:
              service: uhc-acct-mngr
              severity: {{{severity_rhit_errors}}}
            exp_annotations:
              message: "AMS is seeing a high rate of errors from RHIT for 20% of requests."
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/uhc/sop"
      - eval_time: 10m
        alertname: OCM dependencies - Terms and Conditions too many errors - {{{environment}}}
        exp_alerts:
      - eval_time: 20m
        alertname: OCM dependencies - Terms and Conditions too many errors - {{{environment}}}
        exp_alerts:
          - exp_labels:
              service: uhc-acct-mngr
              severity: {{{severity_terms_errors}}}
            exp_annotations:
              message: "AMS is seeing a high rate of errors from Terms and Conditions for 20% of requests."
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/uhc/sop"
      - eval_time: 10m
        alertname: OCM dependencies - Terms and Conditions slow - {{{environment}}}
        exp_alerts:
      - eval_time: 20m
        alertname: OCM dependencies - Terms and Conditions slow - {{{environment}}}
        exp_alerts:
          - exp_labels:
              service: uhc-acct-mngr
              severity: medium
            exp_annotations:
              message: "AMS is seeing a high average latency accessing Terms and Conditions. The average response time over the last 5 minutes is 5 seconds"
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/uhc/sop"
      - eval_time: 1m
        alertname: OCM Alert Pod Restarted - {{{environment}}}
        exp_alerts:
      - eval_time: 3m
        alertname: OCM Alert Pod Restarted - {{{environment}}}
        exp_alerts:
          - exp_labels:
              namespace: uhc-{{{environment}}}
              pod: uhc-acct-mngr-7fccd765b-mpqmk
              service: uhc-acct-mngr
              severity: {{{severity_pod_restart}}}
            exp_annotations:
              message: 'Observed pod (uhc-acct-mngr-7fccd765b-mpqmk) in namespace uhc-{{{environment}}} restart(s) (1)'
      - eval_time: 1m
        alertname: OCM Alert Pod Restarting too many times - {{{environment}}}
        exp_alerts:
      - eval_time: 3m
        alertname: OCM Alert Pod Restarting too many times - {{{environment}}}
        exp_alerts:
          - exp_labels:
              namespace: uhc-{{{environment}}}
              pod: uhc-acct-mngr-fetch-telemeter-metrics-1613509500-srnkx
              service: uhc-acct-mngr
              severity: {{{severity_pod_restart_many_times}}}
            exp_annotations:
              message: 'Observed pod (uhc-acct-mngr-fetch-telemeter-metrics-1613509500-srnkx) in namespace uhc-{{{environment}}} restart(s) (2)'
      - eval_time: 10m
        alertname: OCM dependencies - Hydra too many errors - {{{environment}}}
        exp_alerts:
      - eval_time: 20m
        alertname: OCM dependencies - Hydra too many errors - {{{environment}}}
        exp_alerts:
          - exp_labels:
              service: uhc-acct-mngr
              severity: medium
            exp_annotations:
              message: "AMS is seeing a high rate of errors from Hydra for 10% of requests."
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/uhc/sop#hydra"
      - eval_time: 10m
        alertname: OCM dependencies - Hydra slow - {{{environment}}}
        exp_alerts:
      - eval_time: 20m
        alertname: OCM dependencies - Hydra slow - {{{environment}}}
        exp_alerts:
          - exp_labels:
              service: uhc-acct-mngr
              severity: medium
            exp_annotations:
              message: "AMS is seeing a high average latency accessing Hydra. 10% of the requests took more than 10s during the last 10 minutes."
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/uhc/sop#hydra"
      - eval_time: 60m
        alertname: OCM cron job too many errors - {{{environment}}}
        exp_alerts:
      - eval_time: 120m
        alertname: OCM cron job too many errors - {{{environment}}}
        exp_alerts:
          - exp_labels:
              jobname: ocp-usage-reconciler
              namespace: uhc-{{{environment}}}
              service: uhc-acct-mngr
              severity: medium
            exp_annotations:
              message: "AMS is seeing too many errors from ocp-usage-reconciler in namespace uhc-{{{environment}}}."
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/-/blob/master/docs/uhc/sop/README.md#account-manager-cron-job-error"
    # 5xx tests
    # Normal case, traffic comes from both the router and from clusters service
  - interval: 1m
    input_series:
      - series: haproxy_backend_http_responses_total{code="5xx",exported_namespace="uhc-{{{environment}}}",pod="router-default-69c7459f76-6qz58",route="accounts-mgmt"}
        values: 0+0x5 0+1x15
      - series: haproxy_backend_http_responses_total{code="5xx",exported_namespace="uhc-{{{environment}}}",pod="router-default-69c7459f76-bhswd",route="accounts-mgmt"}
        values: 0+0x5 0+1x15
      - series: haproxy_backend_http_responses_total{code="2xx",exported_namespace="uhc-{{{environment}}}",pod="router-default-69c7459f76-6qz58",route="accounts-mgmt"}
        values: 0+6x20
      - series: haproxy_backend_http_responses_total{code="2xx",exported_namespace="uhc-{{{environment}}}",pod="router-default-69c7459f76-bhswd",route="accounts-mgmt"}
        values: 0+2x20
      - series: envoy_cluster_upstream_rq{envoy_cluster_name="accounts_mgmt",envoy_response_code="500",namespace="uhc-{{{environment}}}",service="clusters-service-metrics-envoy"}
        values: 0+0x5 0+1x15
      - series: envoy_cluster_upstream_rq{envoy_cluster_name="accounts_mgmt",envoy_response_code="503",namespace="uhc-{{{environment}}}",service="clusters-service-metrics-envoy"}
        values: 0+0x5 0+1x15
      - series: envoy_cluster_upstream_rq{envoy_cluster_name="accounts_mgmt",envoy_response_code="204",namespace="uhc-{{{environment}}}",service="clusters-service-metrics-envoy"}
        values: 0+1x20
      - series: envoy_cluster_upstream_rq{envoy_cluster_name="accounts_mgmt",envoy_response_code="200",namespace="uhc-{{{environment}}}",service="clusters-service-metrics-envoy"}
        values: 0+7x20
    alert_rule_test:
      - eval_time: 10m
        alertname: UHC account manager 5xx Errors High - {{{environment}}}
        exp_alerts:
      - eval_time: 20m
        alertname: UHC account manager 5xx Errors High - {{{environment}}}
        exp_alerts:
          - exp_labels:
              service: uhc-acct-mngr
              severity: {{{severity_500_errors}}}
              namespace: uhc-{{{environment}}}
            exp_annotations:
              message: "UHC Account Manager is returning code 5xx for 20% of requests."
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/uhc/sop"

    # All 500s come from the internal traffic
  - interval: 1m
    input_series:
      - series: envoy_cluster_upstream_rq{envoy_cluster_name="accounts_mgmt",envoy_response_code="500",namespace="uhc-{{{environment}}}",service="clusters-service-metrics-envoy"}
        values: 0+0x5 0+1x15
      - series: envoy_cluster_upstream_rq{envoy_cluster_name="accounts_mgmt",envoy_response_code="503",namespace="uhc-{{{environment}}}",service="clusters-service-metrics-envoy"}
        values: 0+0x5 0+1x15
      - series: envoy_cluster_upstream_rq{envoy_cluster_name="accounts_mgmt",envoy_response_code="204",namespace="uhc-{{{environment}}}",service="clusters-service-metrics-envoy"}
        values: 0+1x20
      - series: envoy_cluster_upstream_rq{envoy_cluster_name="accounts_mgmt",envoy_response_code="200",namespace="uhc-{{{environment}}}",service="clusters-service-metrics-envoy"}
        values: 0+7x20
    alert_rule_test:
      - eval_time: 10m
        alertname: UHC account manager 5xx Errors High - {{{environment}}}
        exp_alerts:
      - eval_time: 20m
        alertname: UHC account manager 5xx Errors High - {{{environment}}}
        exp_alerts:
          - exp_labels:
              service: uhc-acct-mngr
              severity: {{{severity_500_errors}}}
              namespace: uhc-{{{environment}}}
            exp_annotations:
              message: "UHC Account Manager is returning code 5xx for 20% of requests."
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/uhc/sop"
   # All 500s come from the router, e.g. pods down
  - interval: 1m
    input_series:
      - series: haproxy_backend_http_responses_total{code="5xx",exported_namespace="uhc-{{{environment}}}",pod="router-default-69c7459f76-6qz58",route="accounts-mgmt"}
        values: 0+0x5 0+1x15
      - series: haproxy_backend_http_responses_total{code="5xx",exported_namespace="uhc-{{{environment}}}",pod="router-default-69c7459f76-bhswd",route="accounts-mgmt"}
        values: 0+0x5 0+1x15
      - series: haproxy_backend_http_responses_total{code="2xx",exported_namespace="uhc-{{{environment}}}",pod="router-default-69c7459f76-6qz58",route="accounts-mgmt"}
        values: 0+6x20
      - series: haproxy_backend_http_responses_total{code="2xx",exported_namespace="uhc-{{{environment}}}",pod="router-default-69c7459f76-bhswd",route="accounts-mgmt"}
        values: 0+2x20
    alert_rule_test:
      - eval_time: 10m
        alertname: UHC account manager 5xx Errors High - {{{environment}}}
        exp_alerts:
      - eval_time: 20m
        alertname: UHC account manager 5xx Errors High - {{{environment}}}
        exp_alerts:
          - exp_labels:
              service: uhc-acct-mngr
              severity: {{{severity_500_errors}}}
              namespace: uhc-{{{environment}}}
            exp_annotations:
              message: "UHC Account Manager is returning code 5xx for 20% of requests."
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/uhc/sop"
  # TODO - interval: 5m
  #   input_series:
  #     # subscription labels
  #     - series: api_inbound_request_duration_sum{namespace="uhc-production",service="uhc-acct-mngr-metrics",path="/api/accounts_mgmt/v1/subscriptions",method="GET",code="200"}
  #       values:
  #     - series: api_inbound_request_duration_count{namespace="uhc-production",service="uhc-acct-mngr-metrics",path="/api/accounts_mgmt/v1/subscriptions",method="GET",code="200"}
  #       values:
  #   promql_expr_test:
  #      - expr: max(subscriptions_request_duration_sum:z_score_prediction)
  #        eval_time: 1d
  #        exp_samples:
  #          - labels: 'subscriptions_request_duration_sum:z_score'
  #            value: 2
  #   alert_rule_test:
  #     - eval_time: 1d
  #       alertname: Z-Score based Duration of requests for Subscriptions are outside of expected range - {{{environment}}}
  #       exp_alerts:
  #         - exp_labels:
  #             service: uhc-acct-mngr
  #             severity: warning
  #           exp_annotations:
  #             runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/uhc/sop"
  #             message: "UHC Account Manager in namespace {{{ environment }}}: Duration of requests for Subscriptions are outside of expected operating parameters"
